## “How Old is Your Dog in Human Years?” Calculator

In this assignment, you will write a program that calculates a dog’s age in human years.
The program will prompt the user for an age in dog years and calculate that age in human years.
Allow for int or float values, but check the user’s input to make sure it's valid -- it should be numeric and positive.
Otherwise, let the user know their input is not valid.

You can use the following rules to approximately convert a medium-sized dog’s age to human years:
- For the first year, one dog year is equal to 15 human years
- For the first 2 years, each dog year is equal to 12 human years
- For the first 3 years, each dog year is equal to 9.3 human years
- For the first 4 years, each dog year is equal to 8 human years
- For the first 5 years, each dog year is equal to 7.2 human years
- After that, each dog year is equal to 7 human years.  (Note: This means the first 5 dog years are
equal to 36 human years (5 * 7.2) and the remaining dog years are equal to 7 human years each.)

Print the result in the following format, substituting for **dog_age** and **human_age**: "The given dog age **dog_age**
is **human_age** in human years."  Round the result to 2 decimal places.  Note: If there is a 0 in the hundredths place,
you can drop it, e.g. 24.00 can be displayed as 24.0.

Considering invalid inputs:
* Your program must ask the user for an age in dog years - hint: use the input() function
* We are going to test invalid inputs - make sure that your code can handle negative value inputs and non-numerical inputs!
* For invalid inputs, make sure that your printed response adheres to the following:
 - If a text-based input is provided, make sure your response contains the word 'invalid'.  
 - If a negative input is provided, make sure your response contains the word 'negative'.  

<!-- grade_id: dog, points: 0, grade: false, solution: true -->
```python
#### import traceback

def calculator():
    
    # Get dog age
    age = float(input("Input dog years: "))

    try:
        # Cast to float
        d_age = age 
       
        # If user enters negative number, print message
        if (age < 0):
            print('Your input: ' + str(age) + ' is negative')
       
        # Otherwise, calculate dog's age in human years
        # your code here
       
        elif(age <= 1.0):
            d_age = 15
        elif(age <= 2): 
            d_age = age * 12
        elif(age <= 3):
            d_age = age * 9.3
        elif(age <= 4):
            d_age = age * 8
        elif(age <= 5):
            d_age = age * 7.2
        elif (age > 5.00):
            d_age = (5 * 7.2) + ((age - 5) * 7)
            
        d_age = round(d_age, 2)
        
        print("The given dog age {} is {} in human years.".format(age, d_age))
        
    except:
        print(age, "is an invalid age.")
        print(traceback.format_exc())
    
calculator()
```

Output:
```
Input dog years: 1
The given dog age 1.0 is 15 in human years.
```

<!-- grade_id: init_test, points: 0, grade: true, solution: false -->
```python
###########################################################
##### EXECUTE THIS CELL BEFORE YOU TEST YOUR SOLUTIONS ####
###########################################################
!pip install nose
import os, sys, re
import unittest
from unittest.mock import patch
from unittest import TestCase
from nose.tools import assert_equal, assert_true, assert_almost_equal
from importlib import reload
from io import StringIO

class RedirectStdStreams(object):
    def __init__(self, stdout=None, stderr=None):
        self._stdout = stdout or sys.stdout
        self._stderr = stderr or sys.stderr

    def __enter__(self):
        self.old_stdout, self.old_stderr = sys.stdout, sys.stderr
        self.old_stdout.flush(); self.old_stderr.flush()
        sys.stdout, sys.stderr = self._stdout, self._stderr

    def __exit__(self, exc_type, exc_value, traceback):
        self._stdout.flush(); self._stderr.flush()
        sys.stdout = self.old_stdout
        sys.stderr = self.old_stderr
        
def parse_dog_age(msg):
    human_age = re.findall("\d+\.\d+ in human", msg)
    human_age = human_age[0].split(' ')[0].strip()
    return float(human_age)
```

Output:
```
Requirement already satisfied: nose in /opt/conda/lib/python3.7/site-packages (1.3.7)
[33mWARNING: You are using pip version 20.2.3; however, version 21.0.1 is available.
You should consider upgrading via the '/opt/conda/bin/python3 -m pip install --upgrade pip' command.[0m
```

<!-- grade_id: test_negative, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

@patch('builtins.input', return_value='-1')
def test_negative(self, mock_input):
    f = StringIO()
    with RedirectStdStreams(stdout=f, stderr=f):
        calculator()
    return str(f.getvalue())

print("Received the below output from your program:\n---------------\n'{}'".format(test_negative(-1)))    
assert_true('negative' in test_negative(-1), "Make sure negative inputs print out a message containing the word negative!")
```

Output:
```
Received the below output from your program:
---------------
'Your input: -1.0 is negative
The given dog age -1.0 is -1.0 in human years.
'
```

<!-- grade_id: test_invalid, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

@patch('builtins.input', return_value='n')
def test_text_input(self, mock_input):
    f = StringIO()
    with RedirectStdStreams(stdout=f, stderr=f):
        calculator()
    return str(f.getvalue())

print("Received the below output from your program:\n---------------\n'{}'".format(test_text_input('n')))    
assert_true('invalid' in test_text_input('n'), "Make sure text inputs print out a message containing the word invalid!")
```

Output:
```
<ipython-input-24-c193cb033d9b> in test_text_input(self, mock_input)
      7     f = StringIO()
      8     with RedirectStdStreams(stdout=f, stderr=f):
----> 9         calculator()
     10     return str(f.getvalue())
     11 

<ipython-input-18-f4d4790fc9c4> in calculator()
      4 
      5     # Get dog age
----> 6     age = float(input("Input dog years: "))
      7 
      8     try:

ValueError: could not convert string to float: 'n'
```

<!-- grade_id: test_1, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

@patch('builtins.input', return_value='1')
def test_one(self, mock_input):
    f = StringIO()
    with RedirectStdStreams(stdout=f, stderr=f):
        calculator()
    return str(f.getvalue())
    

print("Received the below output from your program:\n---------------\n'{}'".format(test_one(1)))    
assert_almost_equal(
    parse_dog_age(test_one(1)), 
    15, 
    "Expected a 1 year old dog to be 15 in human years. Make sure your print statement is formatted as:\n\t \
    'The given dog age <input_value> is <calculated_age> in human years.'")


```

<!-- grade_id: test_2, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

@patch('builtins.input', return_value='2')
def test_two(self, mock_input):
    f = StringIO()
    with RedirectStdStreams(stdout=f, stderr=f):
        calculator()
    return str(f.getvalue())

print("Received the below output from your program:\n---------------\n'{}'".format(test_two(2)))    
assert_almost_equal(parse_dog_age(test_two(2)), 24, 
    "Expected a 2 year old dog to be 24 in human years. Make sure your print statement is formatted as:\n\t \
    'The given dog age <input_value> is <calculated_age> in human years.'")
```

Output:
```
Received the below output from your program:
---------------
'The given dog age 2.0 is 24.0 in human years.
'
```

<!-- grade_id: test_3, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

@patch('builtins.input', return_value='3')
def test_three(self, mock_input):
    f = StringIO()
    with RedirectStdStreams(stdout=f, stderr=f):
        calculator()
    return str(f.getvalue())
    
print("Received the below output from your program:\n---------------\n'{}'".format(test_three(3)))    
assert_almost_equal(parse_dog_age(test_three(3)), 27.9,
            "Expected a 3 year old dog to be 27.9 in human years. Make sure your print statement is formatted as:\n\t \
    'The given dog age <input_value> is <calculated_age> in human years.'")
```

Output:
```
Received the below output from your program:
---------------
'The given dog age 3.0 is 27.9 in human years.
'
```

<!-- grade_id: test_4, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

@patch('builtins.input', return_value='4')
def test_four(self, mock_input):
    f = StringIO()
    with RedirectStdStreams(stdout=f, stderr=f):
        calculator()
    return str(f.getvalue())

print("Received the below output from your program:\n---------------\n'{}'".format(test_four(4)))    
assert_almost_equal(parse_dog_age(test_four(4)), 32,
            "Expected a 4 year old dog to be 32.0 in human years. Make sure your print statement is formatted as:\n\t \
    'The given dog age <input_value> is <calculated_age> in human years.'")
```

Output:
```
Received the below output from your program:
---------------
'The given dog age 4.0 is 32.0 in human years.
'
```

<!-- grade_id: test_5, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

@patch('builtins.input', return_value='5')
def test_five(self, mock_input):
    f = StringIO()
    with RedirectStdStreams(stdout=f, stderr=f):
        calculator()
    return str(f.getvalue())
    
print("Received the below output from your program:\n---------------\n'{}'".format(test_five(5)))    
assert_almost_equal(parse_dog_age(test_five(5)), 36,
            "Expected a 5 year old dog to be 36.0 in human years. Make sure your print statement is formatted as:\n\t \
    'The given dog age <input_value> is <calculated_age> in human years.'")
```

Output:
```
Received the below output from your program:
---------------
'The given dog age 5.0 is 36.0 in human years.
'
```

<!-- grade_id: test_6, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

@patch('builtins.input', return_value='6')
def test_six(self, mock_input):
    f = StringIO()
    with RedirectStdStreams(stdout=f, stderr=f):
        calculator()
    return str(f.getvalue())

print("Received the below output from your program:\n---------------\n'{}'".format(test_six(6)))    
assert_almost_equal(parse_dog_age(test_six(6)), 43,
            "Expected a 6 year old dog to be 43.0 in human years. Make sure your print statement is formatted as:\n\t \
    'The given dog age <input_value> is <calculated_age> in human years.'")
```

Output:
```
Received the below output from your program:
---------------
'The given dog age 6.0 is 43.0 in human years.
'
```

<!-- grade_id: test_7, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

@patch('builtins.input', return_value='7')
def test_seven(self, mock_input):
    f = StringIO()
    with RedirectStdStreams(stdout=f, stderr=f):
        calculator()
    return str(f.getvalue())

print("Received the below output from your program:\n---------------\n'{}'".format(test_seven(7)))
assert_almost_equal(parse_dog_age(test_seven(7)), 50.0,
            "Expected a 7 year old dog to be 50.0 in human years. Make sure your print statement is formatted as:\n\t \
    'The given dog age <input_value> is <calculated_age> in human years.'")
```

Output:
```
Received the below output from your program:
---------------
'The given dog age 7.0 is 50.0 in human years.
'
```

```python

```